# <div style="background-color:rgba(204, 229, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(0, 76, 153, 1);">PHYS 231 Lab #4</span></div>

### Due Monday, Oct. 17 at 2:00 pm $-$ upload your .zip submission to the PHYS 231 Canvas Gradebook.

***
## Lab Learning Outcomes:
By the end of the PHYS 231 Lab #4, students will be able to:
* <b><span style="color:rgba(0, 153, 76, 1);">Design, build, and test practical $LRC$ circuits.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Determine the time constant from the transient response of an $LRC$ circuit.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Perform a weighted linear fit to a set of data and use the extracted best-fit parameters to determine a quantity of interest.</span></b>

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 1 - Installing & Importing Packages</span></div>

## The 'InstallerCheck()' Function:

Run the 'InstallerCheck()' function contained within PHYS231.py by placing you're cursor inside the code cell below and then hitting 'Shift'+'Enter'.  If the function reports that some packages have been installed, run the cell with "InstallerCheck()" a second time.  It should report that all required packages are already installed.

```python
import PHYS231
PHYS231.InstallerCheck()
```

## Importing Packages/Modules:

Once the packages are installed, they can be imported using the notation:
```python
import packageName
```
where 'packageName' is the name of the package to be imported.  Execute the import statements below by placing your cursor within the code cell and hitting 'Shift' + 'Enter'.

```python
# Import required and commonly-used modules
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd
import uncertainties
from uncertainties import ufloat
    
# Initialize Otter
import otter
grader = otter.Notebook("PHYS 231 - Lab 4a.ipynb")
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 2 - Oscilloscope Measurement of the Time Constant</span></div>

## NumPy Arrays

Data that you collect during the lab can be entered directly into a list using the following notation:

```python
x = [x0, x1, x2, x3, ..., xN]
y = [y0, y1, y2, y3, ..., yN]
```
These lists will contain $N+1$ elements. It's usually more convenient to enter data into a NumPy array rather than a list because you can do mathematical operations on arrays that you can't do on lists.  For example:

- squaring an array creates an new array with each of the elements from the original array squared
- multiplying two arrays of the same length creates a new array in which the each element is the product of the corresponding elements from the original arrays
- ...

NumPy arrays are entered using the following syntax:

```python
x = np.array([x0, x1, x2, x3, ..., xN])
y = np.array([y0, y1, y2, y3, ..., yN])
```             

In the cell below, replace the "..." in NumPy array statements with your measurements of time, the peak values of the damped voltage oscillations, and the uncertainty in the peak voltages. 

After executing the cell, the lengths of each of the arrays you've entered will be printed.  Make sure that all of the lengths are the same.

```python
# Enter your data into the arrays below.
# Please do NOT change the variable names.
time = np.array([...]) # s
Vp = np.array([...]) # V"
errVp = np.array([...]) # V

# The statements below will print the lengths of the arrays.
print("Length of time array:", len(time))
print("Length of Vp array:", len(Vp))
print("Length of errVp array:", len(errVp))
```

## Pandas DateFrames

The Pandas package can be used to display your array data in a nicely-formated table.  The Pandas package creates "DataFrames" which are 2 dimensional data structures, like a 2 dimensional array, or a table with rows and columns.  The cell below creates and then displays a DataFrame using the arrays that you created in the cell above.

```python
# Create the data structure.
data = {
    "time (s)": time,
    "VR (V)": Vp,
    "errVR (V)": errVp
        }

# An option to format the numbers in the DataFrame.
pd.set_option('display.float_format', lambda x: '%0.3g' % x)

#Display the data.
pd.DataFrame(data)
```

### Calculate the Natural Logarithm of the Peak Voltage

The transient response of voltage a across are resistor in an highly-underdamped $LRC$ circuit is given by:

\begin{align}
V_\mathrm{R}(t)\approx\omega_1 V_0 RCe^{-\left(\gamma/2\right)t}\sin\left(\omega_1 t\right)
\end{align}

where $V_0$ is the size of the voltage step, $\gamma = R/L$, and $\omega_1=\sqrt{\omega_0^2 - \left(\gamma/2\right)^2}$.  Note that, for the highly-underdamped circuit that we're working with $\omega\gg\gamma/2$ such that $\omega_1\approx\omega_0=1/\sqrt{LC}$. In this case, the peak voltages as a function of time are governed by:

\begin{align}
V_\mathrm{p}(t)\approx\omega_0 V_0 RCe^{-\left(\gamma/2\right)t}
\end{align}

This expression can be linearized by taking the natural logarithm of both sides such that:

\begin{align}
\ln\left(V_\mathrm{p}\right)\approx\left(-\frac{\gamma}{2}\right) t + \ln\left(\omega_0 V_0 RC\right)
\end{align}

This is the equation of a straight line with slope $-\gamma/2$ and $y$-intercept $\ln\left(\omega_0 V_0 RC\right)$.

In the cell below, calculate the natural logarithm of the peak voltage.  Recall that in Python, the natural logarithm of $x$ is called using
```python
np.log(x)
```

```python
lnVp = ...
```

```python
grader.check("ln_voltage")
```

### Calculate the uncertainty in $\ln\left(V_\mathrm{p}\right)$

Recall propagation of errors.  If $y = f\left(x1, x2, ..., x_N\right)$ and $x_i\pm\Delta x_i$ for $i= 1...N$ are known, then:

\begin{align}
\Delta y = \sqrt{\left(\frac{\partial f}{\partial x_1}\Delta x_1\right)^2 + \left(\frac{\partial f}{\partial x_2}\Delta x_2\right)^2 + ... + \left(\frac{\partial f}{\partial x_N}\Delta x_N\right)^2}
\end{align}

If $y=\ln\left(V_\mathrm{p}\right)$, then $y$ is a function of only a single variable and:

\begin{align}
\Delta y = \left\vert \frac{\partial y}{\partial V_\mathrm{p}}\Delta V_\mathrm{p}\right\vert
\end{align}

In the cell below, calculate the uncertainty in $\ln V_\mathrm{p}$.

```python
errlnVp = ...
```

```python
grader.check("ln_uncertainty")
```

### Plot the data

The $\ln\left(V_\mathrm{p}\right)$ data can now be plotted as a function of time.  Execute the cell below to created a scatter plot with error bars.  For additional information about the 'Scatter' function, see **Appendix A** near the end of this notebook. 

```python
fig1 = PHYS231.Scatter(
    xData = time,
    yData = lnVp,
    yErrors = errlnVp,
    xlabel = "time",
    ylabel = "ln Vp",
    xUnits = "s",
    yUnits = "",
)
```

<!-- BEGIN QUESTION -->

## Linear Fits

The final thing that we'll do in Part 2 is fit a straight line through the data corrected (i.e. cut) data.  We'll use the 'LinearFit' contained in PHYS231.py to complete this task.  The function call to perform a weighted fit is:
```python
m, b, dm, db, fig = PHYS231.LinearFit(t, y, erry)
```
A weighted fit with a formatted plot can be called by using some of the other available arguments:
```python
m, b, dm, db, fig = PHYS231.LinearFit(xData = t, yData = y, yErrors = erry, xlabel = 'time', ylabel = 'position', xUnits = 's', yUnits = 'm')
```
Execute the cell below to preform a weighted linear fit to your data.

For additional information about the 'LinearFit' function, see **Appendix B** near the end of this notebook. 

```python
m, b, dm, db, fig2 = PHYS231.LinearFit(
    xData = time,
    yData= lnVp,
    yErrors = errlnVp,
    xlabel = "time",
    ylabel = "ln Vp",
    xUnits = "s",
    yUnits = "",
)
```

<!-- END QUESTION -->

## Determining the time constant: $\gamma\pm\Delta\gamma$

Use the slope from the linear fit and its uncertainty to determine an experimental value for $\gamma=R/L$ and its uncertainty.

```python
# Enter an expression to calculate gamma. 
gamma = ...
```

```python
# Enter an expression to calculate the uncertainty in gamma.
errgamma = ...
```

```python
# The 'uncertainties' package can be used round the numerical value of gamma based on the value of the uncertainty.
# gamma and errgamma are passed to 'ufloat'
print('gamma =', ufloat(gamma, errgamma), '1/s')
```

In your lab notebook, you should compare this experimentally-measured value of $\gamma$ to the value that you expected based on the values of the components used in your circuit.  Recall that $\gamma=R/L$.  Keep in mind that your expected value of $\gamma$ will also have an uncertainty based on the uncertainties in the values of $R\pm\Delta R$ and $L\pm\Delta L$.

<!-- BEGIN QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 3 - Feedback and Submission</span></div>

***
**<span style="color:blue">Question 3.1:</span>**  

We welcome your feedback on the PHYS 231 labs!  Please feel free to include any comments you have about this lab in the cell below.  Your comments will be taken into consideration when revising/improving the PHYS 231 labs.  You can suggest improvements, point out anything that was unclear, comment on the strengths and weaknesses of the lab, ...

This question is optional and will have no impact on your lab grade.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

***
Once you've completed this notebook:
- Save your work.
- Run 'grader.check_all()' to confirm that you've completed all required tasks.
- Run 'grader.export()' to generate a .zip file containing all of the materials that you will submit.
- Download the generated .zip file.
- Upload the .zip file to the PHYS 231 Canvas gradebook.

Here is a <a href = "https://cmps-people.ok.ubc.ca/jbobowsk/phys231/Python/images/Submission.gif">GIF</a> showing how these steps are completed.  Once your completed notebook has been uploaded to the Canvas gradebook, you're done!

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 4 - Playground (optional)</span></div>

Feel free to add as many cells as you like below and use them as a playground for further independent investigations.  These cells won't be graded, so feel free to use them in any way that you like.  

```python
# Here's an empty code cell that you can use.
```

```python
# Here's another empty code cell that you can use.
```

```python
# Here's yet another empty code cell that you can use.
# If you need more, you can add cells using the '+' icon in the menu bar at to the top of the screen.
```

### <div style="background-color:rgba(255, 255, 102, 0.5); text-align:left; padding:20px 0; margin-top:20px">$\quad$Appendix A &ndash; The Scatter Function...</div>

The function for generating scatter plots is called as follows:
```python
PHYS231.Scatter(xData, yData, yErrors = [], xlabel = 'x-axis', ylabel = 'y-axis', xUnits = '', yUnits = '', fill = False, show = True)
```
The 'xData' and 'yData' inputs are required, all other arguments are optional with default values set.  The function returns the a single output (the formatted plot):
```python
fig
```

The function will do a simple scatter plot if no 'yError' are included.  It will included error bars if 'yErrors' are passed to the function.  The 'fill' and 'show' arguments should generally be false.  

If the 'xData' list is empty, as in:
```python
xData = []
```
then the x-axis will be the trial number.  For example, is 'xData' is empty and 'yData' has 10 elements, the x-axis will span 1 to 10.

### Scatter Example Implmentation
The code block below shows an implementation of 'Scatter'.
```python
import PHYS231
theta = [10, 20, 30] # degrees
T = [2.02, 1.95, 2.13] # s
errT = [0.02, 0.02, 0.03] # s
fig = PHYS231.Scatter(theta, T, errT, 'initial angle' , 'period', 'degrees', 's')
```

If you're interested in generating your own plots with customized formatting, see the following Python-based plotting tutorial: https://cmps-people.ok.ubc.ca/jbobowsk/Python/html/Jupyter%20Basic%20Plots.html.
***

### <div style="background-color:rgba(255, 255, 102, 0.5); text-align:left; padding:20px 0; margin-top:20px">$\quad$Appendix B &ndash; The LinearFit Function...</div>

The linear-fit function is called as follows:
```python
PHYS231.LinearFit(xData, yData, yErrors = [], xlabel = 'x-axis', ylabel = 'y-axis', xUnits = '', yUnits = '', fill = False)
```
The xData and yData lists are required.  The others are optional with default values set.  The function returns the following outputs:
```python
Slope, Yintercept, errSlope, errYintercept, fig
```
The first four outputs are numerical values and 'fig' is the formatted plot. 

If $y$-uncertainities are provided, the function will perform a weighted fit.  The $y$-uncertainties list must be the same length as $x$- and $y$-data lists.  If $y$-uncertainties are not provided, the fit will be unweighted.  The other optional arguments include $x$- and $y$-axis names and units for the $x$- and $y$-datasets.  These must be entered as strings (enclosed in quotations) and they are used for formatting the outputs of the function.  If 'fill' is set to True, shading will be added around the best-fit line representing confindence intervals. 

### Unweighted Fit Example Implmentation
The code block below shows the most basic use of 'LinearFit' for an unweighted fit without any of the optional arguments.
```python
import PHYS231
V = [1, 2, 3, 4]
I = [0.12, 0.198, 0.285, 0.412]
m, b, dm, db, fig = PHYS231.LinearFit(V, I)
```

### Weighted Fit Example Implmentation
This second blcok of code shows how to use 'LinearFit' to do a weighted fit.  It also makes use of all of the other optional arguments.  
```python
import PHYS231
V = [1, 2, 3, 4]
I = [0.12, 0.198, 0.285, 0.412]
errI = [0.005, 0.012, 0.020, 0.025]
m, b, dm, db, fig = PHYS231.LinearFit(V, I, errI, 'voltage' , 'current', 'V', 'A', True)
```

The 'm, b, dm, db, fig = ...' syntax is used to store the function outputs in varibles.  If you're interested in details of the fitting process, see the following Python-based tutorial: https://cmps-people.ok.ubc.ca/jbobowsk/Python/html/Jupyter%20Weighted%20Linear%20Fit.html
***

Last update: October 14, 2022

Tests for gamma (points: 1)


Tests for gamma_uncertainty (points: 1)


Tests for ln_uncertainty (points: 1)
```
>>> len(errlnVp) != 0
True
```
```
>>> test = True
>>> for i in range(len(errlnVp)):
...     if np.isclose(errlnVp[i], errVp[i]/Vp[i], atol = 0, rtol = 0.03) != True:
...         test = False
>>> test
True
```


Tests for ln_voltage (points: 1)
```
>>> len(lnVp) != 0
True
```
```
>>> test = True
>>> for i in range(len(lnVp)):
...     if np.isclose(lnVp[i], np.log(Vp[i]), atol = 0, rtol = 0.03) != True:
...         test = False
>>> test
True
```
